---
title: 'Visualization with ggplot2'
description: ""
---

## Introduction to ggplot2

```yaml
type: NormalExercise
key: d05e04f242
xp: 100
```

When installing a packages, you should use the install.packages function the package name should be in quotes

`@instructions`


`@hint`


`@pre_exercise_code`

```{r}
`@pre_exercise_code`

```{r}
load(url("https://assets.datacamp.com/production/repositories/3522/datasets/15037db2827169bf67b3db91ec70cd9e235737dd/Prices_Data12.csv"))
```

```


`@sample_code`

```{r}
___(ggplot2)




ggplot()
```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Scatterplots

```yaml
type: NormalExercise
key: 25e1fbb73c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Visualizing correlations between 2 plots

```yaml
type: NormalExercise
key: 58ccf786be
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
