---
title: 'Part 2 (Cleaning Data)'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: 4cce4d1122
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Insert exercise title here

```yaml
type: NormalExercise
key: ca7514f462
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Insert exercise title here

```yaml
type: NormalExercise
key: dc792567da
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Extra Practice

```yaml
type: NormalExercise
key: dc64edf7dd
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
