---
title: 'Part 2 (Exploring Data)'
description: ""
---

## Exploring your data

```yaml
type: NormalExercise
key: b211902f49
xp: 100
```

glimpse, tibble?

`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Summaries and aggregate functions

```yaml
type: NormalExercise
key: 7754e06b3f
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Visualization for exploration

```yaml
type: NormalExercise
key: 9b1c39a5fd
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Central tendencies and variance

```yaml
type: NormalExercise
key: 751bf1b083
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Extra Practice

```yaml
type: NormalExercise
key: d09b831b37
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
