---
title: 'Introduction to R'
description: ""
---

## An exercise title written in sentence case

```yaml
type: NormalExercise
key: 82a96a0a9b
lang: r
xp: 100
skills: 1
```

You can use the Console located at the bottom to tell R what to do. You can do math, save things as variables, and use functions to perform actions.

`@instructions`
- Copy and paste the first line of code into the Console
- store your output into a variable
- Use the sqrt() function on your variable.

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
sqrt(5+5)

___ <- 5+5

___(___)
```

`@solution`
```{r}
sqrt(5+5)

num <- 5+5

sqrt(num)
```

`@sct`
```{r}
# Update this to something more informative.
success_msg("Good work! You can use functions which work with numbers and variables.")
```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: dde3d6001e
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
