---
title: 'Part 2 (importing with R)'
description: ""
---

## Reading csvs

```yaml
type: NormalExercise
key: 739329f2e8
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Reading data from the Web

```yaml
type: NormalExercise
key: d4c1683124
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Read_xl

```yaml
type: NormalExercise
key: 935f9a910a
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## data.table

```yaml
type: NormalExercise
key: ea6588dc6c
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```


---

## Extra Practice

```yaml
type: NormalExercise
key: a298c62245
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{r}

```


`@sample_code`

```{r}

```


`@solution`

```{r}

```


`@sct`

```{r}

```
